# <p style="text-align: center;"> TP noté R2.07 <br> Quelques algorithmes sur les graphes
</p>

## <span style="color:red">  **Consignes générales :** </span>
1. Lire attentivement le sujet
2. les structures de données ainsi que  les entrées et sorties des méthodes demandées (et leur type) sont **imposées**. Il ne faut pas les modifier.
3. Pour qu'une cellule de code soit prise en compte elle doit êtr
4. e executée.
5. Si vous faites une modification dans une cellule, il faut "l'executer" pour qu'elle soit prise en compte.
6. Des tests de validation de certaines méthodes sont proposés, il est nécessaire de les passer mais pas suffisant !
7. Vous pouvez utiliser, en les adaptant,  les  méthodes obtenues durant les séances de tp. 


```python
from collections import deque
import random
import copy
g1 = {"A" :{"C"},
           "B" : {"C", "E"},
           "C" : {"A", "B", "D", "E"},
           "D" : {"C"},
           "E" : {"C", "B"},
           "F" : set()
         }
g2 ={"A" : set()}
g3 = {
"A":{"B":3,"C":4, "D":7},
"B":{"C":1,"F":5},
"C":{"F":6,"D":2},
"D":{"E":3, "G":6},
"E":{"G":3, "H":4},
"F":{"E":1, "H":8},
"G":{"H":2},
"H":{"G":2}
}
g4 = {
    "A": {"B" :3},
    "B": {"C":1, "F":4},
    "C": {"D":6, "E":2},
    "D": dict(),
    "E": {"F":4},
    "F": dict()
}
```

```python
class Graphe(object):
    
    def __init__(self, graphe_dict=None):
        """ initialise un objet graphe.
	    Si aucun dictionnaire n'est
	    créé ou donné, on en utilisera un 
	    vide
        """
        if graphe_dict == None:
            graphe_dict = dict()
        self._graphe_dict = graphe_dict

    def aretes(self, sommet):
        """ retourne une liste de toutes les aretes d'un sommet"""
        return self._graphe_dict[sommet]

    def all_sommets(self):
        """ retourne tous les sommets du graphe """
        return set(self._graphe_dict.keys())

    def all_aretes(self):
        """ retourne toutes les aretes du graphe """
        return self.__list_aretes()

    def add_sommet(self, sommet):
        """ Si le "sommet" n'set pas déjà présent
	dans le graphe, on rajoute au dictionnaire 
	une clé "sommet" avec une liste vide pour valeur. 
	Sinon on ne fait rien.
        """
        if sommet not in self._graphe_dict:
            self._graphe_dict[sommet] = []

    def add_arete(self, arete):
        """ l'arete est de  type set, tuple ou list;
            Entre deux sommets il peut y avoir plus
	    d'une arete (multi-graphe)
        """
        arete = set(arete)
        arete1, arete2 = tuple(arete)
        for x, y in [(arete1, arete2), (arete2, arete1)]:
            if x in self._graphe_dict:
                self._graphe_dict[x].add(y)
            else:
                self._graphe_dict[x] = {y}

    def __list_aretes(self):
        """ Methode privée pour récupérer les aretes. 
	    Une arete est un ensemble (set)
            avec un (boucle) ou deux sommets.
        """
        aretes = []
        for sommet in self._graphe_dict:
            for voisin in self._graphe_dict[sommet]:
                if ({voisin, sommet})  not in aretes:
                    aretes.append({sommet, voisin})
        return aretes
    
    def __iter__(self):
        self._iter_obj = iter(self._graphe_dict)
        return self._iter_obj

    def __next__(self):
        """ Pour itérer sur les sommets du graphe """
        return next(self._iter_obj)

    def __str__(self):
        res = "sommets: "
        for k in self._graphe_dict.keys():
            res += str(k) + " "
        res += "\naretes: "
        for arete in self.__list_aretes():
            res += str(arete) + " "
        return res
```

## Ecrire la méthode estcyclique, décrite dans l'algorithme ci dessous, qui permet de vérifier si un graphe est **acyclique**.


![Cycle](AlgoCyclique.png)

Si le sommet passé en paramètre n'existe pas, on renvoie une exception. On utlisera la méthode "raise" et  ValueError().

<!-- grade_id: cell-9700bf7fa8adb695, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe):
    """ s est un sommet quelconque, donné par l'utilisateur 
    ou choisi aléatoirement parmi la liste des sommets du graphe
    """
    def estcyclique(self, s=None, TAG=[]):
        if s not in self._graphe_dict: 
            raise ValueError(f"Le sommet {s} n'existe pas dans le graphe.")
        TAG.append(s) # on ajoute le sommet s à la liste TAG
        for voisin in self._graphe_dict[s]: 
            if voisin in TAG: # si le voisin est déjà dans TAG, il y a un cycle
                return "Il y a un cycle"
            if self.estcyclique(voisin, TAG.copy()) == "Il y a un cycle": # sinon on appelle la fonction récursivement
                return "Il y a un cycle"
        return "Pas de cycle"
```

```python
"""Vérifiez vos résultats sur les tests suivants"""
G=Graphe2(g1)
assert G.estcyclique("A") == "Il y a un cycle"
assert G.estcyclique("E") == "Il y a un cycle"

```

<!-- grade_id: cell-065a64a25b12eb9c, points: 1, grade: true, solution: false -->
```python
"""Vérifiez les entrées"""
try:
    G.estcyclique("Z") 
except ValueError:
    pass
else:
    raise AssertionError("n'a pas retourné d'erreur")
```

## Ecrire la méthode estconnexe, décrite ci-dessous, et qui permet de tester si un graphe est connexe

![Connexe](grapheconnexe.png)

<!-- grade_id: cell-4248824b00cbc543, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def estconnexe(self,s):
        """ s est un sommet quelconque, donné par l'utilisateur 
        Il faudra renvoyer une exception si le sommet n'existe pas
        """
        if s not in self._graphe_dict:
            raise ValueError(f"Le sommet {s} n'existe pas dans le graphe.")
        
        visite = set()  # Ensemble des sommets visités
        a_visiter = [s]  # Liste des sommets à visiter
        
        while a_visiter: # Tant qu'il reste des sommets à visiter
            sommet = a_visiter.pop()
            visite.add(sommet) # On ajoute le sommet visité à l'ensemble des sommets visités
            for voisin in self._graphe_dict[sommet]: # On parcourt les voisins du sommet
                if voisin not in visite:
                    a_visiter.append(voisin) # On ajoute le voisin à la liste des sommets à visiter
        
        if len(visite) < len(self.all_sommets()): # Si le nombre de sommets visités est inférieur au nombre de sommets du graphe
            return "faux" 
```

<!-- grade_id: cell-911ea2a527ab4f07, points: 1, grade: true, solution: false -->
```python
"""Vérifiez vos résultats sur les tests suivants"""
G=Graphe2(g1)
assert G.estconnexe("A") == "faux"
assert G.estconnexe("E") == "faux"

```

## Modifier la méthode **estconnexe** pour qu'elle retourne "vrai" si le graphe est connexe et,  sinon, le tuple ("faux",n) où n est le nombre de composantes connexes du graphe.

<!-- grade_id: cell-2e03fc6b59e99f40, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def estconnexe2(self, s):
        """ s est un sommet quelconque, donné par l'utilisateur 
        Il faudra renvoyer une exception si le sommet n'existe pas
        """
        if s not in self._graphe_dict:
            raise ValueError(f"Le sommet {s} n'existe pas dans le graphe.")
        
        visite = set()  # Ensemble des sommets visités
        a_visiter = [s]  # Liste des sommets à visiter
        composantes_connexes = 0  # Compteur de composantes connexes
        
        while a_visiter: # Parcours en profondeur
            sommet = a_visiter.pop()
            visite.add(sommet)
            for voisin in self._graphe_dict[sommet]: # On parcourt les voisins du sommet
                if voisin not in visite:
                    a_visiter.append(voisin) # On ajoute le voisin à la liste des sommets à visiter
        
        composantes_connexes += 1 # On a trouvé une composante connexe
        
        # Parcourir les sommets non visités pour trouver d'autres composantes connexes
        for sommet in self.all_sommets():
            if sommet not in visite: # Si le sommet n'a pas été visité, on a une nouvelle composante connexe
                composantes_connexes += 1
                a_visiter = [sommet]
                while a_visiter: # Parcours en profondeur
                    sommet = a_visiter.pop()
                    visite.add(sommet)
                    for voisin in self._graphe_dict[sommet]: # On parcourt les voisins du sommet
                        if voisin not in visite:
                            a_visiter.append(voisin)
        
        if composantes_connexes > 1: # Si on a plus d'une composante connexe alors le graphe n'est pas connexe
            return ("faux", composantes_connexes)
        else:
            return "vrai"
```

<!-- grade_id: cell-dfe93da380f98b32, points: 0, grade: true, solution: false -->
```python
"""Vérifiez vos résultats sur les tests suivants"""
G=Graphe2(g1)
assert G.estconnexe2("A") == ("faux",2)
assert G.estconnexe2("C") == ("faux",2)

```

## Modifier la méthode précédente pour qu"elle pour qu'elle retourne "vrai" si le graphe est connexe et,  sinon, <br> le tuple (("faux",n), L) où n est le nombre de composantes connexes du graphe et L une liste des composantes connexes (une liste de listes donc).

```python
class Graphe2(Graphe2):
    def estconnexe3(self,s):
        """ s est un sommet quelconque, donné par l'utilisateur 
        Il faudra renvoyer une exception si le sommet n'existe pas
        L=[[]]
        """
        if s not in self._graphe_dict:
            raise ValueError(f"Le sommet {s} n'existe pas dans le graphe.")
        
        visited = set() # Ensemble des sommets visités
        composantes = []

        def dfs(sommet, comp): # Parcours en profondeur
            if sommet not in visited:
                visited.add(sommet)
                comp.append(sommet)
                for voisin in self._graphe_dict[sommet]: # On parcourt les voisins du sommet
                    dfs(voisin, comp)

        for sommet in sorted(self._graphe_dict.keys()): # On parcourt les sommets du graphe
            if sommet not in visited:
                comp = []
                dfs(sommet, comp)
                composantes.append(sorted(comp)) # On ajoute la composante connexe à la liste des composantes connexes
        
        if len(composantes) > 1: # Si on a plus d'une composante connexe alors le graphe n'est pas connexe
            return (("faux", len(composantes)), composantes)
        else:
            return "vrai"
```

```python
"""Vérifiez vos résultats sur les tests suivants"""
G=Graphe2(g1)
assert G.estconnexe3("A") == (("faux",2), [["A","B","C","D","E"],["F"]])
assert G.estconnexe3("C") == (("faux",2), [["A","B","C","D","E"],["F"]])
```

## Dikjstra, plus long chemin

On souhaite ecrire une méthode qui renvoie le plus long chemin à partir d'un sommet donné.<br>
Pour cela on adaptera l'algorithme de Dikjstra pour des graphes  particuliers : **les graphes orientés acycliques pondérés positivement** 

## Ecrire, dans un premier temps, la méthode ci-dessous qui met en oeuvre l'algorithme de Dijkstra classique (D est le tableau des distances et P le dictionnaire des parents, comme vu en tp).

<!-- grade_id: cell-b992084d0ba8f7f3, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def dijkstra(self,s):
        if s not in self._graphe_dict:
            raise ValueError(f"Le sommet {s} n'existe pas dans le graphe.")
        
        D = {sommet: float('infinity') for sommet in self._graphe_dict} # Initialisation des distances, infini pour tous les sommets
        D[s] = 0
        P = {sommet: None for sommet in self._graphe_dict}
        
        sommets = set(self._graphe_dict.keys())
        while sommets: # Tant qu'il reste des sommets à visiter
            u = min(sommets, key=lambda sommet: D[sommet]) # On récupère le sommet avec la distance minimale
            sommets.remove(u) # On retire le sommet de la liste des sommets à visiter
            for voisin in self._graphe_dict[u]:
                if voisin in sommets:
                    alt = D[u] + self._graphe_dict[u][voisin]
                    if alt < D[voisin]:
                        D[voisin] = alt # On met à jour la distance si elle reste infini alors on a pas encore visité le sommet
                        P[voisin] = u      
        return D,P
```

## Ecire la méthode ci dessous qui permet de donner le plus court chemin entre deux sommets, s et d, donnés.<br>Le chemin sera donné sous forme d'une liste dans le bon ordre.

<!-- grade_id: cell-0d85dbb2a97e06ea, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def chemindijkstra(self, s, d):
        """ Cette méthode retourne "pas de chemin" ou un tuple contenant
        le plus court chemin entre s et d sous forme d'une liste, L, dans le bon ordre
        ainsi que la longueur,l, du chemin trouvé. Le chemin d'un sommet 
        à lui-même est de longueur nulle.
        Comme précédemment, on gère une exception si s ou d n'appartiennent pas au graphe.
        """

        if s not in self._graphe_dict or d not in self._graphe_dict:
            raise ValueError("Les sommets doivent appartenir au graphe.")

        D, P = self.dijkstra(s)  # Obtention des distances et des prédécesseurs à partir de l'algorithme de Dijkstra
        if D[d] == float('infinity'):  # Si la destination est inaccessible depuis la source ou non visitée
            return ("pas de chemin", float('infinity'))

        # Reconstruction du chemin
        chemin = []
        sommet_actuel = d
        while sommet_actuel is not None:
            chemin.insert(0, sommet_actuel)
            sommet_actuel = P[sommet_actuel]
        return chemin, D[d]
    
    def dijkstra(self,s):
        if s not in self._graphe_dict:
            raise ValueError(f"Le sommet {s} n'existe pas dans le graphe.")
        
        D = {sommet: float('infinity') for sommet in self._graphe_dict} # Initialisation des distances, infini pour tous les sommets
        D[s] = 0
        P = {sommet: None for sommet in self._graphe_dict}
        
        sommets = set(self._graphe_dict.keys())
        while sommets: # Tant qu'il reste des sommets à visiter
            u = min(sommets, key=lambda sommet: D[sommet]) # On récupère le sommet avec la distance minimale
            sommets.remove(u) # On retire le sommet de la liste des sommets à visiter
            for voisin in self._graphe_dict[u]:
                if voisin in sommets:
                    alt = D[u] + self._graphe_dict[u][voisin]
                    if alt < D[voisin]:
                        D[voisin] = alt # On met à jour la distance si elle reste infini alors on a pas encore visité le sommet
                        P[voisin] = u      
        return D,P
```

<!-- grade_id: cell-c530f17ffb5400d9, points: 3, grade: true, solution: false -->
```python
"""Vérifiez vos résultats sur les tests suivants"""
G3=Graphe2(g3)
assert G3.chemindijkstra("A","E") == (["A","C","D","E"],9)
assert G3.chemindijkstra("E","E") == (["E"],0)

```

On souhaite dans cette partie adapter l'algorithme de Dijkstra pour trouver le plus long chemin entre deux sommets d'un graphe valué orienté.
Pour cela :
1. Si le graphe est acyclique et pondéré positivement sinon on retourne "impossible"
2. Sinon, on utilise l'algorithme de Dijkstra sur le graphe mais en changeant toutes les pondérations en leurs opposées.
3. On retournera alors deux tuples : $(L_1,l_1)$ et $(L_2,l_2)$ où $L_1, L_2, l_1$ et $l_2$ sont respectivement les listes correspondant  au plus court, plus long chemin  (dans le bon ordre) et les longueurs de ces deux chemins
4. Comme précédemment si un des sommets n'existe pas on générera une exception. 

<!-- grade_id: cell-ce7d928e5bce5ed1, points: 0, grade: false, solution: true -->
```python
class Graphe2(Graphe2):
    def cheminsPLC(self,s,d):
        """ Cette méthode retourne "impossible" (si il y a un cycle ou si il y a une pondération négative),
        "pas de chemin" ou deux tuples contenant
        le plus court chemin entre s et d sous forme d'une liste, L1, dans le bon ordre
        ainsi que la longueur,l1, du chemin trouvé, le plus long chemin entre s et d sous forme d'une liste, 
        L2, dans le bon ordre ainsi que la longueur,l2, du chemin trouvé, Le chemin d'un sommet 
        à lui mêeme est de longueur nulle
        Comme précédemment, on gèrera une exception si s ou d n'appartiennent pas au graphe 
        """
        if s not in self._graphe_dict or d not in self._graphe_dict:
            raise ValueError("Les sommets doivent appartenir au graphe.")
        
        if self.estcyclique(s) == "Il y a un cycle":
            return "impossible"
        
        for sommet in self._graphe_dict: # On vérifie si il y a une pondération négative
            for voisin in self._graphe_dict[sommet]: # On parcourt les voisins du sommet
                if self._graphe_dict[sommet][voisin] < 0:
                    return "impossible"
                
        graphe = copy.deepcopy(self._graphe_dict) # Copie du graphe pour ne pas modifier l'original
        for sommet in graphe:
            for voisin in graphe[sommet]:
                graphe[sommet][voisin] *= -1 # On inverse les pondérations

        D1, P1 = self.dijkstra(s) # On applique l'algorithme de Dijkstra sur le graphe original et sur le graphe inversé
        D2, P2 = self.dijkstra(s)
        for sommet in D2: # On inverse les distances
            D2[sommet] *= -1

        # Reconstruction du chemin
        chemin1 = []
        sommet_actuel = d
        while sommet_actuel is not None: # On reconstruit le chemin pour le graphe original
            chemin1.insert(0, sommet_actuel)
            sommet_actuel = P1[sommet_actuel]

        chemin2 = []
        sommet_actuel = d
        while sommet_actuel is not None: # On reconstruit le chemin pour le graphe inversé
            chemin2.insert(0, sommet_actuel)
            sommet_actuel = P2[sommet_actuel]

        l1 = D1[d]
        l2 = D2[d]
        return (chemin1, l1), (chemin2, l2) # On retourne les deux chemins et leurs longueurs

    def dijkstra(self,s):
        if s not in self._graphe_dict:
            raise ValueError(f"Le sommet {s} n'existe pas dans le graphe.")
        
        D = {sommet: float('infinity') for sommet in self._graphe_dict} # Initialisation des distances, infini pour tous les sommets
        D[s] = 0
        P = {sommet: None for sommet in self._graphe_dict}
        
        sommets = set(self._graphe_dict.keys())
        while sommets: # Tant qu'il reste des sommets à visiter
            u = min(sommets, key=lambda sommet: D[sommet]) # On récupère le sommet avec la distance minimale
            sommets.remove(u) # On retire le sommet de la liste des sommets à visiter
            for voisin in self._graphe_dict[u]:
                if voisin in sommets:
                    alt = D[u] + self._graphe_dict[u][voisin]
                    if alt < D[voisin]:
                        D[voisin] = alt # On met à jour la distance si elle reste infini alors on a pas encore visité le sommet
                        P[voisin] = u      
        return D,P
    
    def estcyclique(self, s=None, TAG=[]):
        if s not in self._graphe_dict: 
            raise ValueError(f"Le sommet {s} n'existe pas dans le graphe.")
        TAG.append(s) # on ajoute le sommet s à la liste TAG
        for voisin in self._graphe_dict[s]: 
            if voisin in TAG: # si le voisin est déjà dans TAG, il y a un cycle
                return "Il y a un cycle"
            if self.estcyclique(voisin, TAG.copy()) == "Il y a un cycle": # sinon on appelle la fonction récursivement
                return "Il y a un cycle"
        return "Pas de cycle"
```